Asymmetric cortical projections to striatal direct and indirect pathways distinctly control actions

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

The study by Klug et al. investigated the pathway specificity of corticostriatal projections, focusing on two cortical regions. Using a G-deleted rabies system in D1-Cre and A2a-Cre mice to retrogradely deliver channelrhodopsin to cortical inputs, the authors found that M1 and MCC inputs to direct and indirect pathway spiny projection neurons (SPNs) are both partially segregated and asymmetrically overlapping. In general, corticostriatal inputs that target indirect pathway SPNs are likely to also target direct pathway SPNs, while inputs targeting direct pathway SPNs are less likely to also target indirect pathway SPNs. Such asymmetric overlap of corticostriatal inputs has important implications for how the cortex itself may determine striatal output. Indeed, the authors provide behavioral evidence that optogenetic activation of M1 or MCC cortical neurons that send axons to either direct or indirect pathway SPNs can have opposite effects on locomotion and different effects on action sequence execution. The conclusions of this study add to our understanding of how cortical activity may influence striatal output and offer important new clues about basal ganglia function.

The conceptual conclusions of the manuscript are supported by the data, but the details of the magnitude of afferent overlap and causal role of asymmetric corticostriatal inputs on some behavioral outcomes may be a bit overstated given technical limitations of the experiments.

For example, after virally labeling either direct pathway (D1) or indirect pathway (D2) SPNs to optogenetically tag pathway-specific cortical inputs, the authors report that a much larger number of "non-starter" D2-SPNs from D2-SPN labeled mice responded to optogenetic stimulation in slices than "non-starter" D1 SPNs from D1-SPN labeled mice did. Without knowing the relative number of D1 or D2 SPN starters used to label cortical inputs, it is difficult to interpret the exact meaning of the lower number of responsive D2-SPNs in D1 labeled mice (where only ~63% of D1-SPNs themselves respond) compared to the relatively higher number of responsive D1-SPNs (and D2-SPNs) in D2 labeled mice. While relative differences in connectivity certainly suggest that some amount of asymmetric overlap of inputs exists, differences in infection efficiency and ensuing differences in detection sensitivity in slice experiments make determining the degree of asymmetry problematic.

It is also unclear if retrograde labeling of D1-SPN- vs D2-SPN- targeting afferents labels the same densities of cortical neurons. This gets to the point of specificity in some of the behavioral experiments. If the target-based labeling strategies used to introduce channelrhodopsin into specific SPN afferents label significantly different numbers of cortical neurons, might the difference in the relative numbers of optogenetically activated cortical neurons itself lead to behavioral differences?

We thank the reviewer for the comments and for raising additional interpretations of our results. We agree that determining the relative number of D1- versus D2-SPN starter cells would allow a more accurate estimate of connectivity. However, due to current technical limitations, achieving this level of precision remains challenging. As the reviewer also noted, differences in the number of cortical neurons targeting D1- versus D2-SPNs could introduce additional complexity to the functional effects observed in the behavioral experiments. Moreover, functional heterogeneity is likely to exist not only among cortical neurons projecting to striatal D1- or D2-SPNs, but also within the striatal D1- and D2-SPN populations themselves. Addressing these questions at the single-neuron level will require more refined viral tools in combination with improved recording and manipulation techniques. Despite these limitations, our results suggest that a subpopulation of cortical neurons selectively targets striatal D1-SPNs, supporting a functional dichotomy of pathway-specific corticostriatal subcircuits in the control of behavior.

Reviewer #2 (Public review):

Summary:

Klug et al. use monosynaptic rabies tracing of inputs to D1- vs D2-SPNs in the striatum to study how separate populations of cortical neurons project to D1- and D2-SPNs. They use rabies to express ChR2, then patch D1-or D2-SPNs to measure synaptic input. They report that cortical neurons labeled as D1-SPN-projecting preferentially project to D1-SPNs over D2-SPNs. In contrast, cortical neurons labeled as D2-SPN-projecting project equally to D1- and D2-SPNs. They go on to conduct pathway-specific behavioral stimulation experiments. They compare direct optogenetic stimulation of D1- or D2-SPNs to stimulation of MCC inputs to DMS and M1 inputs to DLS. In three different behavioral assays (open field, intra-cranial self-stimulation, and a fixed ratio 8 task), they show that stimulating MCC or M1 cortical inputs to D1-SPNs is similar to D1-SPN stimulation, but that stimulating MCC or M1 cortical inputs to D2-SPNs does not recapitulate the effects of D2-SPN stimulation (presumably because both D1- and D2-SPNs are being activated by these cortical inputs).

Strengths:

Showing these same effects in three distinct behaviors is strong. Overall, the functional verification of the consequences of the anatomy is very nice to see. It is a good choice to patch only from mCherry-negative non-starter cells in the striatum. This study adds to our understanding of the logic of corticostriatal connections, suggesting a previously unappreciated structure.

Weaknesses:

One limitation is that all inputs to SPNs are expressing ChR2, so they cannot distinguish between different cortical subregions during patching experiments. Their results could arise because the same innervation patterns are repeated in many cortical subregions or because some subregions have preferential D1-SPN input while others do not.

Thank you for raising this thoughtful concern. It is indeed not feasible to restrict ChR2 expression to a specific cortical region using the first-generation rabies-ChR2 system alone. A more refined approach would involve injecting Cre-dependent TVA and RG into the striatum of D1- or A2A-Cre mice, followed by rabies-Flp infection. Subsequently, a Flp-dependent ChR2 virus could be injected into the MCC or M1 to selectively label D1- or D2-projecting cortical neurons. This strategy would allow for more precise targeting and address many of the current limitations.

However, a significant challenge lies in the cytotoxicity associated with rabies virus infection. Neuronal health begins to deteriorate substantially around 10 days post-infection, which provides an insufficient window for robust Flp-dependent ChR2 expression. We have tested several new rabies virus variants with extended survival times (Chatterjee et al., 2018; Jin et al., 2024), but unfortunately, they did not perform effectively or suitably in the corticostriatal systems we examined.

In our experimental design, the aim is to delineate the connectivity probabilities to D1 or D2-SPNs from cortical neurons. Our hypothesis considered includes the possibility that similar innervation patterns could occur across multiple cortical subregions, or that some subregions might show preferential input to D1-SPNs while others do not, or a combination of both scenarios. This leads us to perform a series behavior test that using optogenetic activation of the D1- or D2-projecting cortical populations to see which could be the case.

In the cortical areas we examined, MCC and M1, during behavioral testing, there is consistency with our electrophysiological results. Specifically, when we stimulated the D1-projecting cortical neurons either in MCC or in M1, mice exhibited facilitated local motion in open field test, which is the same to the activation of D1 SPNs in the striatum along (MCC: Fig 3C & D vs. I; M1: Fig 3F & G vs. L). Conversely, stimulation of D2-projecting MCC or M1 cortical neurons resulted in behavioral effects that appeared to combine characteristics of both D1- and D2-SPNs activation in the striatum (MCC: Fig 3C & D vs. J; M1: Fig 3F & G vs. M). The similar results were observed in the ICSS test. Our interpretation of these results is that the activation of D1-projecting neurons in the cortex induces behavior changes akin to D1 neuron activation, while activation of D2-projecting neurons in the cortex leads to a combined effect of both D1 and D2 neuron activation. This suggests that at least some cortical regions, the ones we tested, follow the hypothesis we proposed.

There are also some caveats with respect to the efficacy of rabies tracing. Although they only patch non-starter cells in the striatum, only 63% of D1-SPNs receive input from D1-SPN-projecting cortical neurons. It's hard to say whether this is "high" or "low," but one question is how far from the starter cell region they are patching. Without this spatial indication of where the cells that are being patched are relative to the starter population, it is difficult to interpret if the cells being patched are receiving cortical inputs from the same neurons that are projecting to the starter population. The authors indicate they are patching from mCherry-negative neurons within the region of the mCherry-positive neurons, but since the mCherry population will include both true starter cells and monosynaptically connected cells, this is not perfectly precise. Convergence of cortical inputs onto SPNs may vary with distance from the starter cell region quite dramatically, as other mapping studies of corticostriatal inputs have shown specialized local input regions can be defined based on cortical input patterns (Hintiryan et al., Nat Neurosci, 2016, Hunnicutt et al., eLife 2016, Peters et al., Nature, 2021).

This is a valid concern regarding anatomical studies. Investigating cortico-striatal connectivity at the single-cell level remains technically challenging due to current methodological limitations. At present, we rely on rabies virus-mediated trans-synaptic retrograde tracing to identify D1- or D2-projecting cortical populations. This anatomical approach is coupled with ex vivo slice electrophysiology to assess the functional connectivity between these projection-defined cortical neurons and striatal SPNs. This enables us to quantify connection ratios, for example, the proportion of D1-projecting cortical neurons that functionally synapse onto non-starter D1-SPNs.

To ensure the robustness of our conclusions, it is essential that both the starter cells and the recorded non-starter SPNs receive comparable topographical input from the cortex and other brain regions. Therefore, we carefully designed our experiments so that all recorded cells were located within the injection site, were mCherry-negative (i.e., non-starter cells), and were surrounded by ChR2-mCherry-positive neurons. This configuration ensured that the distance between recorded and starter cells did not exceed 100 µm, maintaining close anatomical proximity and thereby preserving the likelihood of shared cortical innervation within the examined circuitry.

These methodological details are also described in the section on ex vivo brain slice electrophysiology, specifically in the Methods section, lines 453–459:

“D1-SPNs (eGFP-positive in D1-eGFP mice, or eGFP-negative in D2-eGFP mice) or D2-SPNs (eGFP-positive in D2-eGFP mice, or eGFP-negative in D1-eGFP mice) that were ChR2-mCherry-negative, but in the injection site and surrounded by cells expressing ChR2-mCherry were targeted for recording. This configuration ensured that the distance between recorded and starter cells did not exceed 100 µm, maintaining close anatomical proximity and thereby preserving the likelihood of shared cortical innervation within the examined circuitry.”

This experimental strategy was implemented to control for potential spatial biases and to enhance the interpretability of our connectivity measurements.

A caveat for the optogenetic behavioral experiments is that these optogenetic experiments did not include fluorophore-only controls, although a different control (with light delivered in M1) is provided in Supplementary Figure 3. Another point of confusion is that other studies (Cui et al, J Neurosci, 2021) have reported that stimulation of D1-SPNs in DLS inhibits rather than promotes movement. This study may have given different results due to subtly different experimental parameters, including fiber optic placement and NA.

We appreciate the reviewer’s thoughtful evaluation and comments. We have added a short discussion of Cui et al.’s study on optogenetic stimulation of D1-SPNs in the DLS (lines 341-343), which reports findings that contrast with ours and those of other studies.

Reviewer #3 (Public review):

Review of resubmission: The authors provided a response to the reviews from myself and other reviewers. While some points were made satisfactorily, particularly in clarification of the innervation of cortex to striatum and the effects of input stimulation, many of my points remain unaddressed. In several cases, the authors chose to explain their rationale rather than address the issues at hand. A number of these issues (in fact, the majority) could be addressed simply by toning done the confidence in conclusions, so it was disappointing to see that the authors by and large did not do this. I repeat my concerns below and note whether I find them to have been satisfactorily addressed or not.

In the manuscript by Klug and colleagues, the investigators use a rabies virus-based methodology to explore potential differences in connectivity from cortical inputs to the dorsal striatum. They report that the connectivity from cortical inputs onto D1 and D2 MSNs differs in terms of their projections onto the opposing cell type, and use these data to infer that there are differences in cross-talk between cortical cells that project to D1 vs. D2 MSNs. Overall, this manuscript adds to the overall body of work indicating that there are differential functions of different striatal pathways which likely arise at least in part by differences in connectivity that have been difficult to resolve due to difficulty in isolating pathways within striatal connectivity, and several interesting and provocative observations were reported. Several different methodologies are used, with partially convergent results, to support their main points.

However, I have significant technical concerns about the manuscript as presented that make it difficult for me to interpret the results of the experiments. My comments are below.

Major:

There is generally a large caveat to the rabies studies performed here, which is that both TVA and the ChR2-expressing rabies virus have the same fluorophore. It is thus essentially impossible to determine how many starter cells there are, what the efficiency of tracing is, and which part of the striatum is being sampled in any given experiment. This is a major caveat given the spatial topography of the cortico-striatal projections. Furthermore, the authors make a point in the introduction about previous studies not having explored absolute numbers of inputs, yet this is not at all controlled in this study. It could be that their rabies virus simply replicates better in D1-MSNs than D2-MSNs. No quantifications are done, and these possibilities do not appear to have been considered. Without a greater standardization of the rabies experiments across conditions, it is difficult to interpret the results.

This is still an issue. The authors point out why they chose various vectors. I can understand why the authors chose the fluorophores etc. that they did, yet the issues I raised previously are still valid. The discussion should mention that this is a potential issue. It does not necessarily invalidate results, but it is an issue. Furthermore, it is possible (in all systems) that rabies replicates better/more efficiently in some cells than others. This is one possible interpretation that has not really been explored in any study. I don't suggest the authors attempt to do that, but it should be raised as a potential interpretation. If the rabies results could mean several different things, the authors owe it to the readership to state all possible interpretations of data.

We thank the reviewer for the comments and suggestions. Because the same fluorophore (mCherry) was used in both TVA- and ChR2-expressing viruses, it was not possible to distinguish true starter SPNs from TVA-only SPNs or monosynaptically labeled SPNs. This limitation makes it difficult to precisely assess the efficiency of rabies labeling and retrograde tracing in our experimental setup. Moreover, differences in rabies replication efficiency between D1- and D2-SPNs could potentially lead to an apparent lower connection probability from D1-projecting cortical neurons to D2-SPNs than from D2-projecting cortical neurons to D1-SPNs. We have added this clarification to the Discussion (lines 280-297).

The authors claim using a few current clamp optical stimulation experiments that the cortical cells are healthy, but this result was far from comprehensive. For example, membrane resistance, capacitance, general excitability curves, etc are not reported. In Figure S2, some of the conditions look quite different (e.g., S2B, input D2-record D2, the method used yields quite different results that the authors write off as not different). Furthermore, these experiments do not consider the likely sickness and death that occurs in starter cells, as has been reported elsewhere. Health of cells in the circuit is overall a substantial concern that alone could invalidate a large portion, if not all, of the behavioral results. This is a major confound given those neurons are thought to play critical roles in the behaviors being studied. This is a major reason why first-generation rabies viruses have not been used in combination with behavior, but this significant caveat does not appear to have been considered, and controls e.g., uninfected animals, infected with AAV helpers, etc, were not included.

This issue remains unaddressed. I did not request clarity about experimental design, but rather, raised issues about the potential effects of toxicity. I believe this to be a valid concern that needs to be discussed in the manuscript, especially given what look visually like potential differences in S2.

We understand and appreciate the reviewer’s concern regarding the potential cytotoxicity of rabies virus infection. Although we performed the in vivo optogenetic behavioral experiments during a period when rabies-infected cells are generally considered relatively healthy, some deficits in starter cells may still occur and could contribute to the observed effects of optogenetic cortical stimulation. We have added this clarification to the Discussion (lines 298-306).

The overall purity (e.g., EnvA pseudotyping efficiency) of the RABV prep is not shown. If there was a virus that was not well EnvA-pseudotyped and thus could directly infect cortical (or other) inputs, it would degrade specificity. This issue has not been addressed. Viral strain is irrelevant. The quality of the specific preparations used is what matters.

While most of the study focuses on the cortical inputs, in slice recordings, inputs from the thalamus are not considered, yet likely contribute to the observed results. Related to this, in in vivo optogenetic experiments, technically, if the thalamic or other inputs to the dorsal striatum project to the cortex, their method will not only target cortical neurons but also terminals of other excitatory inputs. If this cannot be ruled it, stating that the authors are able to selectively activate the cortical inputs to one or the other population should be toned down.

The authors added text to the discussion to address this point. While it largely does what is intended, based on the one study cited, I disagree with the authors' conclusions that it is "clear" that potential contamination from other sites does not play a role. The simplest interpretation is the one the authors state, and there is some supporting evidence to back up that assertion, but to me that falls short of making the point "clear" that there are no other interpretations.

The statements about specificity of connectivity are not well founded. It may be that in the specific case where they are assessing outside of the area of injections, their conclusions may hold (e.g., excitatory inputs onto D2s have more inputs onto D1s than vice versa). However, how this relates to the actual site of injection is not clear. At face value, if such a connectivity exists, it would suggest that D1-MSNs receive substantially more overall excitatory inputs than D2s. It is thus possible that this observation would not hold over other spatial intervals. This was not explored and thus the conclusions are over-generalized. e.g., the distance from the area of red cells in the striatum to recordings was not quantified, what constituted a high level of cortical labeling was not quantified, etc. Without more rigorous quantification of what was being done, it is difficult to interpret the results.

Again, the goal here would be to make a statement about this in the discussion to clarify limitations of the study. I don't expect the authors to re-do all of these experiments, but since they are discussing the corticostriatal circuits, which have multiple subdomains, this remains a relevant point. It has not been addressed.

The results in Figure 3 are not well controlled. The authors show contrasting effects of optogenetic stimulation of D1-MSNs and D2-MSNs in the DMS and DLS, results which are largely consistent with the canon of basal ganglia function. However, when stimulating cortical inputs, stimulating the inputs from D1-MSNs gives the expected results (increased locomotion) while stimulating putative inputs to D2-MSNs had no effect. This is not the same as showing a decrease in locomotion - showing no effect here is not possible to interpret.

I think that the caveat of showing no clear effects of inputs to D2 stimulation should be pointed out. Yes, I understand that the viruses appeared to express etc., but again it remains possible that the results are driven by a lack of e.g., sufficient ChR2 expression. Aside from a full quantification of the number of cells expressing ChR2, overlap in fiber placement and ChR2 expression (which I don't suggest), this remains a possibility and should be pointed out, as it remains a possibility.

In the light of their circuit model, the result showing that inputs to D2-MSNs drive ICSS is confusing. How can the authors account for the fact that these cells are not locomotor-activating, stimulation of their putative downstream cells (D2-MSNs) does not drive ICSS, yet the cortical inputs drive ICSS? Is the idea that these inputs somehow also drive D1s? If this is the case, how do D2s get activated, if all of the cortical inputs tested net activate D1s and not D2s? Same with the results in Figure 4 - the inputs and putative downstream cells do not have the same effects. Given potential caveats of differences in viral efficiency, spatial location of injections, and cellular toxicity, I cannot interpret these experiments.

The explanation the authors provide in their rebuttal makes sense, however this should be included in the discussion of the manuscript, as it is interesting and relevant.

We thank the reviewer for the valuable comments and suggestions. In line with the reviewer’s recommendation, we have incorporated these explanations into the Discussion (lines 242–279) to help interpret the complex behavioral outcomes of optogenetic stimulation of cortical neurons projecting to D1- or D2-SPNs.

Reviewer #2 (Recommendations for the authors):

I appreciate the authors' responses, which helped clarify some experimental choices. I appreciate that the experiment in Fig S3 serves as a reasonable light control for optogenetics experiments. The careful comparison with methods in Cui et al (2021) is useful, although not added to the main manuscript. Some of the other citations here don't really address the controversy, e.g. Kravitz at al is in DMS, but perhaps fully addressing this issue is outside the scope of the current manuscript and awaits further experiments. I also appreciate the clarification for recording locations that "This configuration ensured that the distance between recorded and starter cells did not exceed 100 µm, maintaining close anatomical proximity and thereby preserving the likelihood of shared cortical innervation within the examined circuitry." However, the statement in the reviewer response does not seem to be added to the manuscript's methods, which I think would be helpful. The criteria for choosing recorded cells are still a bit fuzzy without a map of recording locations and histology. There is also a problem that mCherry-positive cells could be starter cells or could be monosynaptically traced cells, so it is hard to know the area of the starter cell population in these experiments for sure. My evaluation of the manuscript remains largely the same as the original. However, I have adjusted my public review a bit to incorporate the authors' responses. I still think this paper has valuable information, suggesting an interesting and previously unappreciated structure of corticostriatal inputs that I hope this group and others will continue to investigate and incorporate into models of basal ganglia function.

We thank the reviewer for the valuable suggestions. We have now included a comparison with Cui et al. in the Discussion. In addition, we have added the criteria for selecting recorded cells to the Methods section: ‘This configuration ensured that the distance between recorded and starter cells did not exceed 100 µm, maintaining close anatomical proximity and thereby preserving the likelihood of shared cortical innervation within the examined circuitry.’

Changes from the earlier version to the revised version

Abstract

Abstract

The striatal direct and indirect pathways constitute the core for basal ganglia function in action control. Although both striatal D1- and D2-spiny projection neurons (SPNs) receive excitatory inputs from the cerebral cortex, whether or not they share inputs from the same cortical neurons, and how pathway-specific corticostriatal projections control behavior remain largely unknown. Here using a G-deleted rabies system in mice, we found that more than two-thirds of excitatory inputs to D2-SPNs also target D1-SPNs, while only one-third do so vice versa. Optogenetic stimulation of striatal D1- vs. D2-SPN-projecting cortical neurons differently regulate locomotion, reinforcement learning and sequence behavior, implying the functional dichotomy of pathway-specific corticostriatal subcircuits. These results reveal the partially segregated yet asymmetrically overlapping cortical projections on striatal D1- vs. D2-SPNs, and that the pathway-specific corticostriatal subcircuits distinctly control behavior. It has important implications in a wide range of neurological and psychiatric diseases affecting cortico-basal ganglia circuitry.

Introduction

The corticostriatal circuits are critically involved in sensory, cognition, and the learning and control of actions (Aoki et al., 2019; Graybiel, 1998; Haber, 2016; Hikosaka et al., 1998; Jin and Costa, 2010; 2015; Kupferschmidt et al., 2017; Stephenson-Jones et al., 2011; Tanji, 2001; Yin and Knowlton, 2006). Dysfunctional corticostriatal circuitry has been implicated in numerous neurological and psychiatric diseases (Shepherd, 2013), including Parkinson’s (Redgrave et al., 2010), autism (Monteiro and Feng, 2017) and obsessive-compulsive disorder (Dalley and Robbins, 2017). The striatal direct and indirect pathways, made up of D1- vs. D2-expressing spiny projection neurons (SPNs) respectively, constitute the core components for basal ganglia functions in relation to action learning and movement control (Albin et al., 1989; DeLong, 1990; Gerfen et al., 1990). Numerous studies have suggested that the two pathways play distinct yet complementary role in controlling actions (Cui et al., 2013; Geddes et al., 2018; Hikosaka et al., 2019; Hikosaka et al., 2000; Jin et al., 2014; Kravitz et al., 2010; Markowitz et al., 2018; Mink, 2003; Tecuapetla et al., 2016). It is well known that D1- and D2-SPNs are spatially intermixed in the striatum and they both receive major excitatory inputs from the cerebral cortex (Bolam et al., 2000; C.R. Gerfen, 2016; Pan et al., 2010). Previous monosynaptic rabies tracing study has revealed that sensory and limbic cortical regions preferably send projections to D1-SPNs, compared to the motor cortical inputs biased toward D2-SPNs (Wall et al., 2013). However, this anatomical analysis was based on relative percentage of various inputs to a certain population, D1 or D2, and does not reflect the absolute number of cortical projections or the crosstalk of projections to both pathways. Furthermore, how the functional distinction between these two pathways is generated in the corticostriatal circuitry, and whether the striatal D1- and D2-SPNs receive the inputs from the same or different group of cortical neurons remain largely unknown. This is mainly due to the lack of appropriate tools to label and manipulate the specific cortical subpopulations projecting to D1- vs. D2-SPNs for functional investigations.

Here using a G-deleted rabies system in mice (Klug et al., 2018; Osakada et al., 2011; Wall et al., 2013), we are able to selectively target and express channelrhodopsin-2 (ChR2) in presynaptic neurons projecting to D1- vs. D2-expressing SPNs. Whole-cell recordings from brain slice reveal that only one-third of the excitatory inputs to D1-SPNs target D2-SPNs, suggesting that many excitatory inputs to D1-SPNs selectively drive the direct pathway. In contrast, a large proportion of excitatory inputs to D2-SPNs send collateral projections to D1-SPNs, implying that excitatory inputs to D2-SPNs control both the indirect and direct pathways. Optogenetic stimulation of D1- vs. D2-SPN-projecting cortical neurons in vivo differently regulate locomotion, reinforcement learning and sequence behavior, in a cell-type and brain-region dependent manner. These results reveal the functional organization of cell-type- and pathway-specific corticostriatal subcircuits, and offer essential insights into how they might control behavior in health and disease.

Results

A modified rabies virus system (Klug et al., 2018; Osakada et al., 2011; Wall et al., 2013) was employed to label and functionally target the specific cortical neurons projecting to striatal D1- versus D2-SPNs. Specifically, D1- or A2a-Cre mice (Gong et al., 2007) were injected with Cre-dependent helper viruses (AAV5/EF1 -Flex-TVA-mCherry, AAV8/CA-Flex-RG) in the dorsal striatum (Klug et al., 2018; Wall et al., 2013) (Fig. 1A-B; see Materials and Methods). Three weeks later, either (EnvA) SAD-ΔG Rabies-GFP or (EnvA) SAD-ΔG Rabies-ChR2-mCherry was injected into the same striatal location to retrogradely infect the presynaptic cortical neurons projecting to D1- or D2-SPNs (Fig 1B). We first injected (EnvA) SAD-ΔG Rabies-GFP in a subgroup of mice to validate the corticostriatal anatomy. In both D1- and A2a-Cre tracing experiments, intensive labeling was found in different cortical regions as expected including the midcingulate cortex (MCC) (van Heukelum et al., 2020; Vogt and Paxinos, 2014) and the primary motor cortex (M1), which targets mainly the dorsal medial and dorsal lateral striatum respectively (Aoki et al., 2019; Bolam et al., 2000; C.R. Gerfen, 2016; Pan et al., 2010; Shepherd, 2013) (Fig. 1C, D). For functional studies, (EnvA) SAD-ΔG Rabies-ChR2-mCherry was utilized to express ChR2 in the presynaptic cortical neurons projecting to D1- or D2-SPNs. To validate the functional expression of ChR2 in the cortex, whole-cell patch clamp recordings were performed from the mCherry-positive layer V pyramidal neurons in M1 around day 7 post rabies injection (Fig. 1E-G; see Materials and Methods). Both the current-voltage relationship revealed by somatic current injections (Fig. 1H) and the spiking activity elicited by blue laser frequency stimulation (Fig. 1I; Fig. S1) confirmed the overall health and the functional expression of ChR2 in the rabies-infected cortical neurons. These results thus demonstrate that we were able to successfully target and functionally express ChR2 in presynaptic cortical neurons projecting to either striatal D1- or D2-SPNs.

Selective labeling and functional expression of ChR2 in specific cortical neurons projecting to striatal D1- vs. D2-SPNs.

Earlier version: (A) Schematic of SAD-ΔG-Rabies-ChR2-mCherry construct with the glycoprotein deleted and replaced with ChR2-mCherry. (B) Timeline of viral injections of Cre-dependent helper viruses and the modified rabies virus for slice and behavioral experiments. (C) Example of coronal brain section with rabies-eGFP injection in the dorsal medial striatum of D1-Cre (top) or A2a-Cre (bottom) mouse shows enriched eGFP expression in the MCC. Scale bar, 1 mm. Inset (right): Higher magnification of retrogradely-labeled striatal D1– or D2-SPN projecting neurons in the MCC expressing eGFP. Dotted lines demarcate cortical lamina. Scale bar, 200 µm. (D) Similar experiments of labeling striatal D1– vs. D2-SPN projecting neurons in M1 with rabies-eGFP. (E) Cartoon brain schematic of ChR2-mCherry expressing M1 neurons projecting to SPNs (red) during whole-cell patch clamp recordings. (F) Example of coronal brain section with rabies-ChR2-mCherry injection in the dorsal lateral striatum of A2a-Cre mouse. Scale bar, 1 mm. Inset (right): Higher magnification of retrogradely-labeled striatal D2-SPN projecting neurons in the M1 showed clear membrane expression of ChR2-mCherry. Scale bar, 200 µm. (G) (left) 10x epifluorescent (red channel) of ChR2-mCherry positive neurons in M1. Scale bar, 250 µm. (middle) 40x image of a patched layer 5 pyramidal neuron under DIC optics. Scale bar, 50 µm. (right) Epifluorescent image (red channel) showing patched layer 5 pyramidal neuron somas expressing ChR2-mCherry signal. Red dotted line denotes patched neuron. Scale bar, 50 µm. (H) Current-voltage traces of a ChR2-mCherry positive layer 5 M1 neuron under current clamp responding to hyperpolarizing and depolarizing current injection steps. Scale bars, 200 ms, 25 mV. (I) Optogenetic stimulation (20 Hz) elicits robust action potentials with high fidelity in a ChR2-mCherry positive D1-SPN projecting M1 neuron in layer 5. Scale bars, 100 ms, 20 mV.

Revised version: (A) Schematic of SAD-ΔG-Rabies-ChR2-mCherry construct with the glycoprotein deleted and replaced with ChR2-mCherry. (B) Timeline of viral injections of Cre-dependent helper viruses and the modified rabies virus for slice and behavioral experiments. (C) Example of coronal brain section with rabies-eGFP injection in the dorsal medial striatum of D1-Cre (top) or A2a-Cre (bottom) mouse shows enriched eGFP expression in the MCC. Scale bar, 1 mm. Inset (right): Higher magnification of retrogradely-labeled striatal D1- or D2-SPN projecting neurons in the MCC expressing eGFP. Dotted lines demarcate cortical lamina. Scale bar, 200 µm. (D) Similar experiments of labeling striatal D1- vs. D2-SPN projecting neurons in M1 with rabies-eGFP. (E) Cartoon brain schematic of ChR2-mCherry expressing M1 neurons projecting to SPNs (red) during whole-cell patch clamp recordings. (F) Example of coronal brain section with rabies-ChR2-mCherry injection in the dorsal lateral striatum of A2a-Cre mouse. Scale bar, 1 mm. Inset (right): Higher magnification of retrogradely-labeled striatal D2- SPN projecting neurons in the M1 showed clear membrane expression of ChR2-mCherry. Scale bar, 200 µm. (G) (left) 10× epifluorescent (red channel) of ChR2-mCherry positive neurons in M1. Scale bar, 250 µm. (middle) 40× image of a patched layer 5 pyramidal neuron under DIC optics. Scale bar, 50 µm. (right) Epifluorescent image (red channel) showing patched layer 5 pyramidal neuron somas expressing ChR2- mCherry signal. Red dotted line denotes patched neuron. Scale bar, 50 µm. (H) Current-voltage traces of a ChR2-mCherry positive layer 5 M1 neuron under current clamp responding to hyperpolarizing and depolarizing current injection steps. Scale bars, 200 ms, 25 mV. (I) Optogenetic stimulation (20 Hz) elicits robust action potentials with high fidelity in a ChR2-mCherry positive D1-SPN projecting M1 neuron in layer 5. Scale bars, 100 ms, 20 mV.

Taking advantage of this rabies-ChR2 system, we first sought to determine how many functional excitatory inputs that the striatal D1- and D2-SPNs might share. The possible functional organization of excitatory inputs to D1- and D2-SPNs at the single cell level, like the corticostriatal projections, could be completely segregated, totally overlapping, or partially mixed (Fig. 2A). In order to differentiate these possibilities, we made whole-cell recordings from D1- or D2-SPNs in brain slice by optogenetic stimulation of rabies-ChR2-infected excitatory terminals in striatum. We asked what the probability is that a D1- or D2-SPN targeted by the same presynaptic excitatory inputs projecting to the nearby D1- or D2-SPN population. D1- or A2a-Cre mice were crossed to the D1- or D2-eGFP reporter line for visualizing striatal D1- vs. D2-SPNs in slice recordings (see Methods). Following the helper viruses and rabies-ChR2-mCherry injection in the D1-/A2a-Cre x D1-/D2-eGFP mice, the mCherry negative striatal SPNs were selected to be recorded in the whole-cell mode and D1- vs. D2-SPNs can be further separated based on the eGFP expression. Picrotoxin, a GABAA antagonist, was added throughout the recordings to isolate the excitatory postsynaptic currents (EPSCs). Following the blue laser stimulation of ChR2-positive presynaptic terminals in striatum, the short-latency (< 10 ms) EPSCs recorded was considered as the direct excitatory inputs on D1- or D2-SPNs (Klug et al., 2018; Kress et al., 2013), which can be blocked by glutamate antagonists NBQX/APV (see Methods).

The excitatory inputs to striatal D1- vs. D2-SPNs are partially segregated with asymmetrical overlapping.

Earlier version: (A) Schematic for the possible organization of the excitatory inputs to striatal D1– vs. D2-SPNs from completely segregated (left), totally overlapping (middle), to partially mixed (right). The red and black filled circles indicate the individual neurons projecting to D1– vs. D2-SPNs, respectively. The half red and half black circles imply the neurons projecting to both. (B) (left) Schematic of rabies-ChR2 labeling of the inputs to D1-SPNs and whole-cell recordings of rabies-negative striatal D1-SPNs with local optogenetic stimulation. (right) Example of the average EPSC trace showing short latency response to paired pulses (50 ms ISI) stimulation (black), that is blocked by AMPAR and NMDAR antagonists (gray). All recordings were conducted in the presence of picrotoxin (PTX) to isolate excitatory transmission. Scale bar, 25 ms, 100 pA. Same conditions applied to all following recordings. (C) Whole-cell recording of rabies-negative striatal D2-SPNs with local optogenetic stimulation with rabies-ChR2 labeling of the inputs to D2-SPNs. (D) The likelihood of the inputs to D1-SPNs form a functional connection with nearby non-starter D1-SPNs, and the likelihood of the D2-SPN situation. Numbers above the bars denote number of cells that show functional connectivity within total recorded. Fisher’s exact test, P = 0.3137. Number of animals in each group: D1-projecting to D1 EGFP(+), N = 7; D1-projecting to D2 EGFP(−), N = 8; D2-projecting to D2 EGFP(+), N = 10; D2-projecting to D1 EGFP(−), N = 8. (E-F) Whole-cell recording of rabies-negative striatal D2-SPNs with local optogenetic stimulation with rabies-ChR2 labeling of the inputs to D1-SPNs (E), and recording of rabies-negative D1-SPNs with stimulation of inputs to D2-SPNs (F). (G) The likelihood of the inputs to D1-SPNs form a functional connection with nearby non-starter D2-SPNs, and the likelihood of the inputs to D2-SPNs form a functional connection with nearby non-starter D1-SPNs. Fisher’s exact test, P = 0.0079. **, P < 0.01. Number of animals in each group: D1-projecting to D2 EGFP(+), N = 8; D1-projecting to D1 EGFP(−), N = 7; D2-projecting to D1 EGFP(+), N = 8; D2-projecting to D2 EGFP(−), N = 10.

Revised version: (A) Schematic for the possible organization of the excitatory inputs to striatal D1- vs. D2-SPNs from completely segregated (left), totally overlapping (middle), to partially mixed (right). The red and black filled circles indicate the individual neurons projecting to D1- vs. D2-SPNs, respectively. The half red and half black circles imply the neurons projecting to both. (B) (left) Schematic of rabies-ChR2 labeling of the inputs to D1-SPNs and whole-cell recordings of rabies-negative striatal D1- SPNs with local optogenetic stimulation. (right) Example of the average EPSC trace showing short latency response to paired pulses (50 ms ISI) stimulation (black), that is blocked by AMPAR and NMDAR antagonists (gray). All recordings were conducted in the presence of picrotoxin (PTX) to isolate excitatory transmission. Scale bar, 25 ms, 100 pA. Same conditions applied to all following recordings. (C) Whole- cell recording of rabies-negative striatal D2-SPNs with local optogenetic stimulation with rabies-ChR2 labeling of the inputs to D2-SPNs. (D) The likelihood of the inputs to D1-SPNs form a functional connection with nearby non-starter D1-SPNs, and the likelihood of the D2-SPN situation. Numbers above the bars denote number of cells that show functional connectivity within total recorded. Fisher’s exact test, P = 0.3137. Number of animals in each group: D1-projecting to D1 EGFP(+), N = 7; D1-projecting to D2 EGFP(-), N = 8; D2-projecting to D2 EGFP(+), N = 10; D2-projecting to D1 EGFP(-), N = 8. (E-F) Whole-cell recording of rabies-negative striatal D2-SPNs with local optogenetic stimulation with rabies- ChR2 labeling of the inputs to D1-SPNs (E), and recording of rabies-negative D1-SPNs with stimulation of inputs to D2-SPNs (F). (G) The likelihood of the inputs to D1-SPNs form a functional connection with nearby non-starter D2-SPNs, and the likelihood of the inputs to D2-SPNs form a functional connectionwith nearby non-starter D1-SPNs. Fisher’s exact test, P = 0.0079. **, P < 0.01. Number of animals in each group: D1-projecting to D2 EGFP(+), N = 8; D1-projecting to D1 EGFP(-), N = 7; D2- projecting to D1 EGFP(+), N = 8; D2-projecting to D2 EGFP(-), N = 10.

Earlier version: Recordings from the mCherry-negative, non-starter striatal D1-SPNs in striatal D1-rabies-ChR2-infected mice revealed that with high probability (∼63%) a D1-SPN receives the inputs from the presynaptic excitatory neurons projecting to surrounding D1-SPNs (Fig. 2B, D; Fig. S2). This is true from recordings in non-starter D1-SPNs identified both as mCherry (−) / eGFP (+) in D1-Cre x D1-eGFP mice and mCherry (−) / eGFP (−) in D1-Cre x D2-eGFP mice (Fig. 2B, D). Similarly, recordings from mCherry-negative non-starter striatal D2-SPNs in striatal D2-rabies-ChR2-tracing mice revealed that with a very high probability (∼79%) a D2-SPN receives the inputs from the presynaptic excitatory neurons project to surrounding D2-SPNs (Fig. 2C, D; Fig. S2). Again, it is similar from recordings in non-starter D2-SPNs identified both as mCherry (−) / eGFP (+) in A2a-Cre x D2-eGFP mice and mCherry (−) / eGFP (−) in A2a-Cre x D1-eGFP mice (Fig. 2C, D). However, recordings from striatal D2-SPNs in the striatal D1-rabies-ChR2-tracing mice revealed that the chance for a D2-SPN to receive excitatory inputs from the presynaptic neurons projecting to surrounding D1-SPNs is rather low (∼40%, Fig. 2E, G; Fig. S2). In contrast, recordings from striatal D1-SPNs in the striatal D2-rabies-ChR2-tracing mice revealed that the chance for a D1-SPN to receive the excitatory inputs from the presynaptic neurons projecting to surrounding D2-SPNs is remarkably high (∼73%, Fig. 2F, G; Fig. S2).

Revised version: Recordings from the mCherry-negative, non-starter striatal D1-SPNs in striatal D1-rabies-ChR2-infected mice revealed that with high probability (∼63%) a D1-SPN receives the inputs from the presynaptic excitatory neurons projecting to surrounding D1-SPNs (Fig. 2B, D; Fig. S2). This is true from recordings in non-starter D1-SPNs identified both as mCherry (-) / eGFP (+) in D1-Cre x D1-eGFP mice and mCherry (-) / eGFP (-) in D1-Cre x D2-eGFP mice (Fig. 2B, D). Similarly, recordings from mCherry-negative non-starter striatal D2-SPNs in striatal D2-rabies-ChR2-tracing mice revealed that with a very high probability (∼79%) a D2-SPN receives the inputs from the presynaptic excitatory neurons project to surrounding D2-SPNs (Fig. 2C, D; Fig. S2). Again, it is similar from recordings in non-starter D2-SPNs identified both as mCherry (-) / eGFP (+) in A2a-Cre x D2-eGFP mice and mCherry (-) / eGFP (-) in A2a-Cre x D1-eGFP mice (Fig. 2C, D). However, recordings from striatal D2-SPNs in the striatal D1-rabies-ChR2-tracing mice revealed that the chance for a D2-SPN to receive excitatory inputs from the presynaptic neurons projecting to surrounding D1-SPNs is rather low (∼40%, Fig. 2E, G; Fig. S2). In contrast, recordings from striatal D1-SPNs in the striatal D2-rabies-ChR2-tracing mice revealed that the chance for a D1-SPN to receive the excitatory inputs from the presynaptic neurons projecting to surrounding D2-SPNs is remarkably high (∼73%, Fig. 2F, G; Fig. S2). These data unveil a complex picture including both parallel and crosstalk between the excitatory inputs to D1- and D2-SPNs. Notably, the likelihood that the input connectivity was significantly higher from the presynaptic excitatory inputs of D2-SPNs to D1-SPNs than from the presynaptic excitatory inputs of D1-SPNs to D2-SPNs (Fig. 2D, G). Together these results suggest largely segregated yet asymmetrically overlapping excitatory projections to striatum where the majority of excitatory inputs to D1-SPNs only target the D1-SPNs, while most excitatory inputs to D2-SPNs target both D2- and D1-SPNs.

Removed in the revised version: These data unveil a complex picture including both parallel and crosstalk between the excitatory inputs to D1– and D2-SPNs. Notably, the likelihood that the input connectivity was significantly higher from the presynaptic excitatory inputs of D2-SPNs to D1-SPNs than from the presynaptic excitatory inputs of D1-SPNs to D2-SPNs (Fig. 2D, G). Together these results suggest largely segregated yet asymmetrically overlapping excitatory projections to striatum where the majority of excitatory inputs to D1-SPNs only target the D1-SPNs, while most excitatory inputs to D2-SPNs target both D2– and D1-SPNs.

Based on this asymmetrically overlapping functional organization, one would predict that the excitatory inputs to D1-SPNs mostly control the striatal direct pathway, while the inputs to D2-SPNs would drive both the indirect and direct pathways (Fig. 3A). To test whether this is the case, we injected rabies-ChR2-mCherry into the dorsal striatum of D1- or A2a-Cre mice as before, and implanted optic fibers bilaterally in either MCC or M1 (see Methods). This allows us to selectively activate D1- or D2-SPN projecting neurons in MCC or M1 and determine the optogenetic effects on behavior. For comparison, we performed behavioral experiments by optogenetic stimulation of striatal D1- or D2-SPNs in dorsal medial (DMS) and dorsal lateral striatum (DLS), two areas that receive dense excitatory projections from MCC and M1, respectively (Aoki et al., 2019; Shepherd, 2013) (see Methods). Consistent with the previous observations (Kravitz et al., 2010), optogenetic stimulation (20Hz) of D1-SPNs in the DMS or DLS facilitated locomotion (Fig. 3B, C, E, F). Conversely, optogenetic stimulation (20Hz) of D2-SPNs in DMS significantly suppressed locomotion (Fig. 3B, D), which is less obvious in DLS (Fig. 3E, G).

Different effects of optogenetic stimulation of D1- vs. D2-SPN projecting cortical neurons on locomotion and reinforcement learning.

(A) Schematic of largely segregated yet partially overlapping excitatory inputs to striatal D1- vs. D2-SPNs. (B) Schematic of dorsal medial striatum (DMS) injection of Cre-dependent AAV-ChR2 and optogenetic simulation in D1- or A2a-Cre mice. (C) (top) Example of locomotion path under control (black) and following 20Hz optogenetic stimulation (gray) of DMS D1-SPNs in open field. Scale bar, 5cm, same for below. (bottom) Stimulation of D1-SPNs in DMS facilitates locomotion (n = 5, unpaired two-tailed t-test, t = 3.386, P = 0.0046). **, P < 0.01. (D) 20Hz stimulation of D2-SPNs in DMS suppresses locomotion (n = 5, unpaired two-tailed t-test, t = 2.559, P = 0.0227). *, P < 0.05. (E) Schematic for dorsal lateral striatum (DLS) optogenetics. (F-G) 20Hz stimulation of D1-SPNs in DLS facilitates locomotion (F, n = 5, unpaired two-tailed t-test, t = 4.736, P = 0.0003), while stimulation of D2-SPNs in DLS does not significantly suppress locomotion in open field (G, n = 5, unpaired two-tailed t-test, t = 1.026, P = 0.3224). ***, P < 0.001. (H) Schematic of rabies-ChR2 labeling of the inputs to D1 or D2-SPNs and optogenetic stimulation in MCC. (I-J) 20Hz stimulation of MCC neurons projecting to D1-SPNs facilitates locomotion (I, n = 9, unpaired two-tailed t-test, t = 2.344, P = 0.0344), while stimulation of MCC neurons projecting to D2-SPNs does not alter locomotion (J, n = 10, unpaired two-tailed t-test, t = 1.214, P = 0.2447). *, P < 0.05. (K) Schematic of rabies-ChR2 labeling of the inputs to D1 or D2-SPNs and optogenetic stimulation in M1. (L-M) 20Hz stimulation of the M1 neurons projecting to D1-SPNs facilitates locomotion (L, n = 7, unpaired two-tailed t-test, t = 3.276, P = 0.0055), while stimulation of the M1 neurons projecting to D2-SPNs does not significantly alter locomotion (M, n = 8, Unpaired two-tailed t-test, t = 0.5796, P = 0.5714). **, P < 0.01. (N) Schematic of a mouse performing intracranial self-stimulation (ICSS) behavior. (O-P) D1-SPN (red) but not D2-SPN stimulation (black) drives ICSS behavior in both the DMS (O: D1, n = 6, permutation test, slope = 1.5060, P = 0.0378; D2, n = 5, permutation test, slope = -0.2214, P = 0.1021; one-tailed Mann Whitney test, Day 7 D1 vs. D2, P = 0.0130) and the DLS (P: D1, n = 6, permutation test, slope = 28.1429, P = 0.0082; D2, n = 5, permutation test, slope = -0.3429, P = 0.0463; one-tailed Mann Whitney test, Day 7 D1 vs. D2, P = 0.0390). *, P < 0.05. (Q) Timeline of helper virus injections, rabies-ChR2 injections and optogenetic stimulation for ICSS behavior. (R-S) Optogenetic stimulation of the cortical neurons projecting to either D1- or D2-SPNs induces ICSS behavior in both the MCC (R: MCC-D1, n = 5, permutation test, Day1- Day7, slope = 2.5857, P = 0.0034; MCC-D2, n = 5, Day2-Day7, permutation test, slope = 1.4229, P = 0.0344; no significant effect on Day7, MCC-D1 vs. MCC-D2, two-tailed Mann Whitney test, P = 0.9999) and the M1 (S: M1-D1, n = 5, permutation test, Day1-Day7, slope = 1.8214, P = 0.0259; M1-D2, n = 5, Day1-Day7, permutation test, slope = 1.8214, P = 0.0025; no significant effect on Day7, M1-D1 vs. M1- D2, two-tailed Mann Whitney test, P = 0.3810). n.s., not statistically significant.

Notably, high-frequency (20Hz) but not low-frequency (5Hz) optogenetic stimulation of MCC neurons that project to D1-SPNs significantly facilitated locomotion in the open field (Fig. 3H, I; Fig. S1), similar to D1-SPN activation in DMS. However, optogenetic stimulation (20Hz) of D2-SPN projecting MCC neurons in the same location did not alter locomotion in the open field (Fig. 3H, J), in contrast with the effects of stimulation of D2-SPNs in DMS (Fig. 3D). Similarly, high-frequency optogenetic stimulation (20Hz) of M1 neurons that project to D1-SPNs facilitated locomotion in the open field (Fig. 3K, L; Fig. S1), while 20Hz stimulation of the M1 neurons projecting to D2-SPNs did not significantly alter locomotion (Fig. 3K, M). Further control experiments employing the same optogenetic stimulation in the exact cortical locations but with ChR2 expression only in the striatum do not generate any behavioral phenotypes (Fig. S3). It thus rules out the possibility that the behavioral effects observed by cortical stimulation in the rabies-ChR2 mice were triggered through direct striatal activation due to the light penetration into the striatum. These results are consistent with the functional connectivity in which the excitatory inputs to D1-SPNs mostly drive the direct pathway, and the inputs to D2-SPNs target both the indirect and direct pathways (Fig. 3A). It also suggests that the cortical neurons in the same cortical layer and spatial location could differently control actions depending on their striatal projection targets, in a pathway- and cell type-specific manner.

We next ask whether the cortical subpopulations projecting to striatal D1- vs. D2-SPNs could differently control action learning. We first performed experiments in the D1-Cre mice with viral expression of ChR2 in the striatum, and found that optogenetic stimulation of D1-SPNs robustly supported intracranial self-stimulation (ICSS) (Fig. 3N) in either DMS (Fig. 3O) or DLS (Fig. 3P). Conversely, optogenetic stimulation of D2-SPNs, either in DMS (Fig. 3O) or DLS (Fig. 3P), did not promote ICSS behavior. These data confirmed that the D1-SPN activation in both DMS and DLS drives action learning and ICSS, while D2-SPN stimulation does not strongly support ICSS behavior (Kravitz et al., 2012; Vicente et al., 2016).

We then test how the striatum-projecting cortical neurons in MCC or M1 would support ICSS behavior, and whether there is any difference between activation of the D1- vs. D2-SPN projecting cortical neurons. Similar to the effects of direct striatal D1-SPN stimulation (Fig. 3O, P), optogenetic stimulation of striatal D1-SPN projecting neurons was sufficient to support ICSS behavior both in MCC (Fig. 3Q, R) and in M1 (Fig. 3Q, S). Notably, optogenetic stimulation of the cortical neurons projecting to D2-SPNs also significantly drove ICSS behavior, irrespective of whether it is in MCC (Fig. 3R) or M1 (Fig. 3S). These data suggested that optogenetic activation of either D1- or D2-SPN projecting neurons in MCC or M1 could drive reinforcement learning and support ICSS behavior.

Corticostriatal circuitry is critical for action sequence learning and execution (Geddes et al., 2018; Hikosaka et al., 1998; Jin and Costa, 2010; 2015; Jin et al., 2014; Tanji, 2001; Tecuapetla et al., 2016). In particular, striatal direct and indirect pathways have been suggested to play distinct roles in controlling learned action sequences, as D1-SPNs facilitate ongoing actions while D2-SPNs inhibit actions and mediate switching (Geddes et al., 2018; Jin and Costa, 2010; 2015; Jin et al., 2014; Tecuapetla et al., 2016). We thus ask how the D1- vs. D2-SPN projecting neurons in MCC and M1 regulate the learned action sequences. D1- or A2a-Cre mice injected with helper viruses were trained under fixed-ratio schedule, in which a fixed amount of eight (FR8) leads to reward (Geddes et al., 2018; Jin and Costa, 2010; Jin et al., 2014; Tecuapetla et al., 2016) (Fig. 4A; see Methods). Three weeks later, the trained animals were injected with (EnvA) SAD-ΔG Rabies-ChR2-mCherry virus in the dorsal striatum and optic fibers were bilaterally implanted in either MCC or M1 as before. Mice were continuously trained for around 4 days after 4 days of recovery from the surgeries to allow the rabies-mediated ChR2 expression before the optogenetic experiments start (Fig. 4E). High-frequency stimulation (20Hz) of the cortical neurons projecting to D1-SPNs or D2-SPNs was delivered upon the first lever press of the FR8 sequence in randomly chosen 50% trials (Geddes et al., 2018; Tecuapetla et al., 2016) (Fig. 4A, E, see Methods). Stimulation of MCC inputs to D1-SPNs facilitated lever pressing over the duration of the FR8 sequence (Fig. 4B, D). Conversely, stimulation of MCC inputs to D2-SPNs slightly reduced the lever press rate over the stimulation period (Fig. 4C, D). The modulation effects on lever pressing rate were significantly different between optogenetic stimulation of D1- and D2-SPN projecting MCC neurons (Fig. 4D). On the other hand, optogenetic activation of the M1 neurons that project to D1-SPNs facilitated lever pressing during sequence execution (Fig. 4F, H), similar to the effects of MCC stimulation. However, optogenetic stimulation of the M1 neurons projecting to D2-SPNs delivered an overall facilitation effect on lever pressing (Fig. 4G, H). Overall, stimulation of either D1- or D2-SPN projecting M1 neurons facilitated lever pressing in a similar degree (Fig. 4H). These results thus revealed the highly heterogeneous functions of corticostriatal subcircuits in controlling learned action sequences, depending on both the cortical region and their cell-type specific targets in striatum.

Optogenetic stimulation of D1- vs. D2-SPN projecting cortical neurons differently modulates action sequence execution.

Earlier version: (A) Schematic of a mouse performing FR8 sequence. (B) Optogenetic stimulation (20Hz) of the D1-SPN projecting MCC neurons during FR8 sequence. Example lever pressing (black bar) in control (top) vs. stimulation (middle) trials aligned to the first press, where the blue transparent rectangle corresponds to the window of optogenetic stimulation (20Hz, 8s). The black and blue lines in the PETH (bottom) indicate the lever press rate for control and stimulation conditions, respectively, same for below. (C) Optogenetic stimulation (20Hz) of the D2-SPN projecting MCC neurons during FR8 sequence. (D) Average percent lever press rate change during optogenetic stimulation of D1– vs. D2-SPN projecting MCC neurons compared to control (MCC – D1, n = 8; MCC – D2, n = 7; Unpaired two-tailed t-test, t = 2.774, P = 0.0097). **, P < 0.01. Effects compared to a theoretical percentage change of 0 of each individual manipulation (MCC-D1, n = 8, one-sample two-tailed t-test, t = 2.814, P = 0.0131; MCC-D2, n = 7, one-sample two-tailed t-test, t = 0.8481, P = 0.4117). (E) Timeline of helper virus injections, rabies-ChR2 injections and optogenetic stimulation during action sequence performance. (F-G) Optogenetic stimulation (20Hz) of the D1– (F) or D2-SPN (G) projecting M1 neurons during FR8 sequence. (H) Average percent lever press rate change during optogenetic stimulation of D1– vs. D2-SPN projecting M1 neurons compared to control (M1 – D1, n = 6; M1 – D2, n = 7; Unpaired two-tailed t-test, t = 0.7651, P = 0.4511). Effects compared to a theoretical percentage change of 0 of each individual manipulation (M1-D1, n = 6, one-sample two-tailed Wilcoxon signed-rank test, P = 0.0046; M1-D2, n = 7, one-sample two-tailed Wilcoxon signed-rank test, P = 0.0479).

Revised version: (A) Schematic of a mouse performing FR8 sequence. (B) Optogenetic stimulation (20Hz) of the D1-SPN projecting MCC neurons during FR8 sequence. Example lever pressing (black bar) in control (top) vs. stimulation (middle) trials aligned to the first press, where the blue transparent rectangle corresponds to the window of optogenetic stimulation (20Hz, 8s). The black and blue lines in the PETH (bottom) indicate the lever press rate for control and stimulation conditions, respectively, same for below. (C) Optogenetic stimulation (20Hz) of the D2-SPN projecting MCC neurons during FR8 sequence. (D) Average percent lever press rate change during optogenetic stimulation of D1- vs. D2-SPN projecting MCC neurons compared to control (MCC – D1, n = 8; MCC – D2, n = 7; Unpaired two-tailed t-test, t = 2.774, P = 0.0097). **, P < 0.01. Effects compared to a theoretical percentage change of 0 of each individual manipulation (MCC-D1, n = 8, one-sample two-tailed t-test, t = 2.814, P = 0.0131, 95% CI, 15.49 to 112.2; MCC-D2, n = 7, one-sample two-tailed t-test, t = 0.8481, P = 0.4117, 95% CI, - 21.78 to 9.502). (E) Timeline of helper virus injections, rabies-ChR2 injections and optogenetic stimulation during action sequence performance. (F-G) Optogenetic stimulation (20Hz) of the D1- (F) or D2-SPN (G) projecting M1 neurons during FR8 sequence. (H) Average percent lever press rate change during optogenetic stimulation of D1- vs. D2-SPN projecting M1 neurons compared to control (M1 – D1, n = 6; M1 – D2, n = 7; Unpaired two-tailed t-test, t = 0.7651, P = 0.4511). Effects compared to a theoretical percentage change of 0 of each individual manipulation (M1-D1, n = 6, one-sample two-tailed Wilcoxon signed-rank test, P = 0.0046, 97.75% CI, -0.7866 to 151.0; M1-D2, n = 7, one-sample two- tailed Wilcoxon signed-rank test, P = 0.0479, 96.48% CI, 2.350 to 62.86).

Discussion

By taking advantage of a monosynaptic rabies tracing with optogenetics system, we have discovered a significant degree of segregation between the excitatory inputs to striatal D1- vs. D2-SPNs. Notably, the results unveiled an overall asymmetric crosstalk from the excitatory inputs of D2-SPNs onto D1-SPNs, but not vice versa. Striatal D1- and D2-SPNs receive excitatory inputs from both the cortex and thalamus (Klug et al., 2018; Wall et al., 2013). Since the current techniques do not allow us to isolate the inputs from a specific region to D1- vs. D2-SPNs in slice recording, these results do not exclude the possibility that there might be certain cortical or thalamic regions targeting D1- and D2-SPNs equally or even with a reverse bias. However, the overall functional organization does imply that while the excitatory inputs to D1-SPNs in general drive the striatal direct pathway, the excitatory inputs to D2-SPNs control both the striatal direct and indirect pathways. Indeed, it has been recently reported that corticospinal neurons, which project to both spinal cord and DLS, form uneven synapses onto direct and indirect pathway neurons in the DLS and preferentially target at D1- other than D2-SPNs (Nelson et al., 2021). Furthermore, a series of in vivo optogenetic experiments in both MCC and M1 have further supported this notion, and demonstrated that the functionally heterogeneous corticostriatal neuronal subpopulations differently control actions, in both a cortical-region- and striatal-targeting-cell-type-specific manner. These in vivo functional findings in corticostriatal pathways are in consistent with the observations of in vitro synapse connection probability. Future studies should aim to further dissect the organization and function of pathway-specific thalamostriatal subcircuits, and determine whether they share the same principles of corticostriatal projections.

Added in the revised version: Our findings show that stimulation of D1-projecting cortical neurons produced behavioral effects closely resembling those of selective D1 activation in both open field and ICSS tests. This is consistent with our slice recording data, which revealed that D1-projecting cortical neurons exhibit a higher connection probability with D1-SPNs than with D2-SPNs. In contrast, interpreting the effects of D2-projecting cortical neuron stimulation is inherently more nuanced. In the open field test, activation of these neurons did not significantly modulate local motion. This could reflect a balanced influence of D1 activation, which facilitates movement, and D2 activation, which suppresses it - resulting in a net neutral behavioral outcome. In the ICSS test, the absence of a strong reinforcement effect typically associated with D2 activation, combined with partial reinforcement likely due to concurrent D1 activation, suggests that stimulation of D2-projecting neurons produces a mixed behavioral signal. This outcome supports the interpretation that these neurons synapse onto both D1- and D2-SPNs, leading to a blended behavioral response that differs from selective D1 or D2 activation alone. Together, these two behavioral assays offer complementary perspectives, providing a more complete view of how projection-specific cortical inputs influence striatal output and behavior.

Added in the revised version: In Figure 4 of the current manuscript, we show that optogenetic activation of MCC neurons projecting to D1-SPNs facilitates sequence lever pressing, whereas activation of MCC neurons projecting to D2-SPNs does not induce significant behavioral changes. Conversely, activation of M1 neurons projecting to either D1- or D2-SPNs enhances lever pressing sequences. These observations align with our prior findings (Geddes et al., 2018; Jin et al., 2014), where we demonstrated that in the striatum, D1-SPN activation facilitates ongoing lever pressing, whereas D2-SPN activation is more involved in suppressing ongoing actions and promoting transitions between sub-sequences. Taken together, the facilitation of lever pressing by D1-projecting MCC and M1 neurons is consistent with their preferential connectivity to D1-SPNs and their established behavioral role.

Added in the revised version: What is particularly intriguing, though admittedly more complex, is the behavioral divergence observed upon activation of D2-SPN-projecting cortical neurons. Activation of D2-projecting MCC neurons does not alter lever pressing, possibly reflecting a counterbalancing effect from concurrent D1- and D2-SPN activation. In contrast, stimulation of D2-projecting M1 neurons facilitates lever pressing, albeit less robustly than their D1-projecting counterparts. This discrepancy may reflect regional differences in striatal targets, DMS for MCC versus DLS for M1, as also supported by our open field test results. Furthermore, our recent findings (Zhang et al., 2025) show that synaptic strength from Cg to D2-SPNs is stronger than to D1-SPNs, whereas the M1 pathway exhibits the opposite pattern. These data suggest that beyond projection ratios, synaptic strength also shapes cortico-striatal functional output. Thus, stronger D2-SPN synapses in the DMS may offset D1-SPN activation during MCC-D2 stimulation, dampening lever pressing increase. Conversely, weaker D2 synapses in the DLS may permit M1-D2 projections to facilitate behavior more readily.

Added in the revised version: To identify non-starter D1- or D2-SPNs, we used mCherry to label both TVA helper virus and ChR2-expressing rabies constructs, allowing us to record mCherry-negative non-starter SPNs labeled by GFP to calculate connection probability. However, this approach introduces a limitation: it is not possible to accurately assess rabies infection efficiency, because both TVA-only SPNs and monosynaptically traced SPNs exhibit mCherry signal in addition to true starter SPNs. Existing data indicate that the connection probabilities from D1-projecting cortical neurons to both D1- and D2-SPNs are lower than those from D2-projecting cortical neurons to both SPN subtypes. One possible explanation is that rabies virus replicates more efficiently in D1-SPNs than in D2-SPNs, or vise versa, causing the calculated connection ratios to reflect a combination of true connectivity and differences in viral replication efficiency. On the other hand, we recorded from non-starter SPNs within the injection area that were surrounded by mCherry-positive starter cells, which should share similar topographic cortical projections.

Added in the revised version: Under these conditions, the number of labeled starter SPNs is unlikely to substantially influence the connection probability measured in brain slices. Furthermore, our recent study (Li and Jin, 2023) proposes a ‘triple-control’ model of striatal function, consisting of a direct pathway operating linearly and two indirect pathway layers serving modulatory roles. This framework suggests a functional requirement for D2-SPNs in the indirect pathway to receive stronger cortical excitatory input than D1-SPNs, which is consistent with the connection data obtained in the present study.

Added in the revised version: Although monosynaptic rabies tracing has been widely applied in anatomical studies, its use in functional and in vivo experiments remains limited by toxicity. In our in vivo optogenetic behavioral experiments, we confined all stimulations to within 10 days post-infection, a period during which rabies-infected cells are reported to remain relatively healthy (Osakada et al., 2011). Nevertheless, toxicity may still compromise neuronal excitability, disrupt synaptic transmission, or even cause the loss of starter SPNs. As a result, the observed optogenetic effects could partly arise from biased circuit function due to such cellular alterations. Addressing this issue will require improved monosynaptic retrograde tracing tools with reduced toxicity and more reliable labeling efficiency.

Earlier version: The cortical neurons projecting to striatum mainly consist of layer 2/3 and layer 5 pyramidal cells (Klug et al., 2018; Wall et al., 2013), including both the intratelencephalic (IT) and pyramidal tract (PT) types of neurons (Shepherd, 2013). While some anatomical preference might exist (Lei et al., 2004), it has been found that both the striatal direct and indirect pathways receive functional inputs from both the IT and PT neurons (Ballion et al., 2008; Kress et al., 2013). Our rabies-ChR2 tracing system allows us to further separate the cortical inputs to striatal D1– vs. D2-SPNs and selectively stimulate these specific cortical subpopulations during behavior and learning. These results have further revealed the diversity of corticostriatal cell subtypes and underscored their heterogeneous functions in behavior. Although the behavioral phenotypes of optogenetic stimulation of different cortical neuronal subpopulations are largely consistent with their functional connectivity with the striatal D1– vs. D2-SPNs, it does not necessarily suggest the observed effects were mediated completely by striatum but not through their collaterals targeting other brain regions or spinal cord (Nelson et al., 2021; Shepherd, 2013). In addition, it has been known that both striatal direct and indirect pathways receive inhibitory inputs from certain GABAergic interneurons in motor cortices (Melzer et al., 2017). In our behavioral experiments with optogenetic stimulation in the motor cortex, there might be possible contribution from these striatum-projecting cortical inhibitory neurons. However, given the nature of sparse distribution of the GABAergic interneurons in the cortex, it is unlikely that they dominate the observed behavioral phenotype (Melzer et al., 2017). Due to the limitations of current viral approaches, one might question whether our in vivo optogenetic experiments could inadvertently activate terminals from rabies-infected neurons in cortical or thalamic regions beyond the targeted area. However, a previous optogenetic study on thalamo-cortical synapses (Cruikshank et al., 2010) demonstrated that sustained optogenetic stimulation (500 ms) of ChR2-expressing presynaptic terminals induces transient, rather than sustained, action potential firing on the postsynaptic neurons. Considering the much longer stimulation durations used in our optogenetic experiments (15 s for open-field, 1 s for ICSS, and 8 s for the FR4/FR8), it is clear that the behavioral effects of optogenetic stimulation are most likely driven by the excitation of ChR2-expressing D1– or D2-SPN projecting neurons in the targeted cortical area, rather than by efferent terminals from neurons located outside this region. Nevertheless, from the striatum point of view, the distinct behavior effect does strongly suggest that the specific information the direct vs. indirect pathway received from the cortex is somehow channeled, but at the same time, effectively coordinated by the cortex.

Revised version: The cortical neurons projecting to striatum mainly consist of layer 2/3 and layer 5 pyramidal cells (Klug et al., 2018; Wall et al., 2013), including both the intratelencephalic (IT) and pyramidal tract (PT) types of neurons (Shepherd, 2013). While some anatomical preference might exist (Lei et al., 2004), it has been found that both the striatal direct and indirect pathways receive functional inputs from both the IT and PT neurons (Ballion et al., 2008; Kress et al., 2013). Our rabies-ChR2 tracing system allows us to further separate the cortical inputs to striatal D1- vs. D2-SPNs and selectively stimulate these specific cortical subpopulations during behavior and learning. These results have further revealed the diversity of corticostriatal cell subtypes and underscored their heterogeneous functions in behavior. Although the behavioral phenotypes of optogenetic stimulation of different cortical neuronal subpopulations are largely consistent with their functional connectivity with the striatal D1- vs. D2-SPNs, it does not necessarily suggest the observed effects were mediated completely by striatum but not through their collaterals targeting other brain regions or spinal cord (Nelson et al., 2021; Shepherd, 2013). In addition, it has been known that both striatal direct and indirect pathways receive inhibitory inputs from certain GABAergic interneurons in motor cortices (Melzer et al., 2017). In our behavioral experiments with optogenetic stimulation in the motor cortex, there might be possible contribution from these striatum-projecting cortical inhibitory neurons. However, given the nature of sparse distribution of the GABAergic interneurons in the cortex, it is unlikely that they dominate the observed behavioral phenotype (Melzer et al., 2017). Due to the limitations of current viral approaches, one might question whether our in vivo optogenetic experiments could inadvertently activate terminals from rabies-infected neurons in cortical or thalamic regions beyond the targeted area. However, a previous optogenetic study on thalamo-cortical synapses (Cruikshank et al., 2010) demonstrated that sustained optogenetic stimulation (500 ms) of ChR2-expressing presynaptic terminals induces transient, rather than sustained, action potential firing on the postsynaptic neurons. Considering the much longer stimulation durations used in our optogenetic experiments (15 s for open-field, 1 s for ICSS, and 8 s for the FR4/FR8), it provides compelling evidence that the behavioral effects of optogenetic stimulation are most likely driven by the excitation of ChR2-expressing D1- or D2-SPN projecting neurons in the targeted cortical area, rather than by efferent terminals from neurons located outside this region. Nevertheless, from the striatum point of view, the distinct behavior effect does strongly suggest that the specific information the direct vs. indirect pathway received from the cortex is somehow channeled, but at the same time, effectively coordinated by the cortex.

Earlier version: These results have important implications on how the corticostriatal circuitry controls actions in health and disease. The traditional model of the basal ganglia suggests that the direct and indirect pathways play opponent roles in facilitating and inhibiting action, respectively (Albin et al., 1989; DeLong, 1990; Kravitz et al., 2010). More recent models of basal ganglia, however, propose that the direct pathway co-activates and cooperates with the indirect pathway with the former activating the selected action and the latter inhibiting the competing actions (Cui et al., 2013; Hikosaka et al., 2000; Jin et al., 2014; Mink, 1996; Tecuapetla et al., 2016). Under more complicated behavior context, it has been previously reported that the striatal D1– and D2-SPNs are co-activated during the initiation of an action sequence, but become largely segregated during the sequence performance (Geddes et al., 2018; Jin et al., 2014). More specifically, the various subpopulations of striatal D1– and D2-SPNs differently change their firing activity to support the start/stop of the sequence, the execution of the elemental actions, and the switch between subsequences (Geddes et al., 2018). These previous findings thus suggested that the striatal direct and indirect pathways have to dynamically coordinate their activity throughout the performance of sequential actions (Geddes et al., 2018; Hikosaka et al., 2019; Jin and Costa, 2015; Markowitz et al., 2018; Tecuapetla et al., 2016).

Revised version: These results have important implications on how the corticostriatal circuitry controls actions in health and disease. The traditional model of the basal ganglia suggests that the direct and indirect pathways play opposing roles in facilitating and inhibiting action, respectively (Albin et al., 1989; DeLong, 1990; Kravitz et al., 2010), although one study (Cui et al., 2021) reported an alternative pattern, which may stem from differences in stimulation coordinates and fiber-optic diameter. More recent models of basal ganglia, however, propose that the direct pathway co-activates and cooperates with the indirect pathway with the former activating the selected action and the latter inhibiting the competing actions (Cui et al., 2013; Hikosaka et al., 2000; Jin et al., 2014; Mink, 1996; Tecuapetla et al., 2016). Under more complicated behavior context, it has been previously reported that the striatal D1- and D2-SPNs are co-activated during the initiation of an action sequence, but become largely segregated during the sequence performance (Geddes et al., 2018; Jin et al., 2014). More specifically, the various subpopulations of striatal D1- and D2-SPNs differently change their firing activity to support the start/stop of the sequence, the execution of the elemental actions, and the switch between subsequences (Geddes et al., 2018). These previous findings thus suggested that the striatal direct and indirect pathways have to dynamically coordinate their activity throughout the performance of sequential actions (Geddes et al., 2018; Hikosaka et al., 2019; Jin and Costa, 2015; Markowitz et al., 2018; Tecuapetla et al., 2016).

But how are the dynamically different activities in the striatal direct and indirect pathways generated in the circuitry? Both the striatal direct and indirect pathways are driven by the excitatory inputs from the cerebral cortex and thalamus (Bolam et al., 2000; C.R. Gerfen, 2016; Pan et al., 2010; Wall et al., 2013). However, whether or not they receive the projections from the same presynaptic neurons, and how the input information is channeled into the two pathways for proper action control remain mostly unknown. The current study has revealed the largely segregated but asymmetrically overlapping organization of the cortical projections to striatal direct vs. indirect pathway. This specific corticostriatal organization provides a structural foundation for the striatal direct and indirect pathways to implement such a dynamic coordination of activity during sequence behavior (Geddes et al., 2018; Hikosaka et al., 2019; Hikosaka et al., 1998; Jin and Costa, 2010; 2015; Jin et al., 2014; Markowitz et al., 2018; Tanji, 2001; Tecuapetla et al., 2016). For instance, the dedicated cortical projections to striatal direct vs. indirect pathway are well suited for controlling sequence initiation and termination, where the activation of D1- and D2-SPNs is critical (DeLong, 1990; Geddes et al., 2018). On the other hand, the overlapping cortical projections to both striatal direct and indirect pathways could be crucial for action switching, which requires proper coordination of the two pathways to inhibit current action and activate the upcoming one (DeLong, 1990; Geddes et al., 2018). Our findings also predict that the striatal D1- vs. D2-SPN projecting neurons in the cerebral cortex would fire differently but activate in relation with each other during behavior. Future work should aim to understand how these two cortical subpopulations behave and coordinate to control the striatal direct and indirect pathways for action learning and selection in health and disease (Dalley and Robbins, 2017; Geddes et al., 2018; Hikosaka et al., 2000; Jin et al., 2014; Mink, 2003; Monteiro and Feng, 2017; Redgrave et al., 2010; Shepherd, 2013).

Materials and Methods

Animals

All procedures were approved by the Salk Institute Institutional Animal Care and Use Committee and followed NIH guidelines for the care and use of laboratory animals. Group housed male and female mice (2 - 6 months old) were used in this study. Animals were housed on a 12-hour dark/12-hour light cycle (dark from 6 pm to 6 am). Heterozygous Drd1-Cre (The Jackson Laboratory, stock # 030329, GENSAT: EY217) and Adora2a-Cre (The Jackson Laboratory, stock # 036158, GENSAT: KG139) mice were obtained from MMRRC and were backcrossed to C57Bl6/J mice, stock # 000664 (> 9 generations) (Cui et al., 2013; Jin et al., 2014; Madisen et al., 2012; Tecuapetla et al., 2016). BAC reporter lines D1-eGFP (MMRRC: MMRRC_000297-MU; GENSAT: X60) and D2-eGFP (MMRRC: MMRRC_00230-UNC; GENSAT: S118) (Gong et al., 2007) were crossed to Drd1-Cre (D1-Cre) and Adora2a-Cre (A2a-Cre) mice to identify D1- and D2-SPNs for electrophysiological recordings.

Surgery and viral injection

For G-deleted rabies-mediated retrograde tracing and functional determination (slice recordings) (Smith et al., 2016), all surgeries were performed under aseptic conditions with animals anesthetized with ketamine (100 mg/kg) / xylazine (10 mg/kg) while mounted on a stereotaxic device (Kopf Instruments; Tujunga, CA). The skull was leveled at bregma and lambda and a small hole was drilled at the coordinate (from bregma and midline) of AP + 0.5 mm, ML ± 1.8mm. A Hamilton syringe (33-gauge needle) containing 1 µl freshly mixed AAV5/EF1 -Flex-TVA-mCherry (UNC Vector Core; Chapel Hill, NC) and AAV8/CA-Flex-RG (UNC Vector Core; Chapel Hill, NC) was slowly lowered to DV - 2.2 mm from the dura to target dorsal central striatum. The virus cocktail was injected slowly over ∼10 min, and the needle was left in place for ∼5 additional minutes afterwards. Then, the needle was slowly retracted over 5 minutes to reduce the virus from moving into the needle track. After injection, mice were sutured and returned to their home cage with ibuprofen (50 mg/kg/day) in their drinking water for the following four days. They were given three weeks to allow for maximal expression of helper viruses, before they were injected with 1.5 µl of (EnvA) SAD-ΔG Rabies-eGFP or 1.5 µl of (EnvA) SAD-ΔG Rabies-ChR2-mCherry (Salk Vector Core, La Jolla, CA) on an angle (18°) to avoid labeling any neurons in the initial injection tract in the same target region. Injecting locations were identical in D1-Cre and A2a-Cre animals. All the injections were done unilaterally for anatomical and slice physiology experiments, and bilaterally for behavioral experiments.

To prepare animals for optogenetic behavior experiments testing D1- or D2-SPN projecting cortical neurons, animals were anesthetized with isoflurane (4% induction, 1-2% maintenance) and locally injected with bupivacaine to numb the incision site. The animals received bilateral injections of helper virus (TVA, RG) as before in dorsal striatum. After ∼21 days of pre-training and full body weight recovered (see Operant Conditioning), the skull was exposed again, and cleaned with 4% H2O2 and UV-light etched with Opti-Bond All-in-One (Kerr, Orange, CA). Then 1.5 µl (EnvA) SAD-ΔG Rabies-ChR2-mCherry was bilaterally injected in each hemisphere using the same coordinates as before. Then, custom made, polished optical fibers (200 µm diameter, 0.37 NA; Thor Labs, Newton, NJ) were implanted in input regions: MCC (AP +0.2 mm, ML ±0.8 mm for skull holes, fibers penetrate into brain at 17° angle off midline with traveling distance of 1.3mm, actual fiber tips target brain at AP +0.2 mm, ML ±0.4 mm, DV -1.2 mm) or M1 ( AP +0.5 mm, ML ±1.2 mm, DV -0.5 mm). The fibers were secured with a light-curing composite (Tetric EvoFlow, Ivoclar Vivadent; Mississauga, ON). Finally, a layer of black dental cement (Lang Dental, Wheeling, IL) was added on the top of the previous cement to support and block laser light diffusion during stimulation. Animals were given ibuprofen in their drinking water for pain management during post-surgery recovery (4 days).

For striatal opto-ICSS and open field experiments, D1- or A2a-Cre mice were injected bilaterally with AAV5-EF1 -DIO-ChR2(H134R)-mCherry (Salk Vector Core, La Jolla, CA) in DMS (AP 0.5 mm, ML ±1.5 mm, DV -2.2 mm) or DLS (AP 0.5 mm, ML ±2.5 mm, DV -2.2 mm), and fiber optics were implanted ∼0.2 mm above the injection site. In control experiments for testing striatal activation by light penetration from cortical optic fibers (Fig S3), D1- or A2a-Cre mice were injected with AAV5-EF1 -DIO-ChR2(H134R)-mCherry bilaterally in DMS, and fiber optics were bilaterally implanted into M1 of the same coordinates as previously described.

Ex vivo brain slice electrophysiology

Earlier version: 4-8 days were allowed for expression and optimal cell health post unilateral (EnvA) SAD-ΔG Rabies-ChR2-mCherry injection before electrophysiology recordings on acute slice were carried out (Klug et al., 2017; Smith et al., 2016). Mice were anesthetized with ketamine/xylazine and transcardially perfused with ∼20 mL ice-cold, bubbling (95% O2/5% CO2) NMDG cutting solution [consisting of (in mM): NMDG 105, HCl 105, KCl 2.5, NaH2PO4 1.2, NaHCO3 26, Glucose 25, Sodium L-Ascorbate 5, Sodium Pyruvate 3, Thiourea 2, MgSO4 10, CaCl2 0.5, 300 mOsm, pH = 7.4]. The extracted brain was blocked coronally with a brain matrix (Zivic Instruments; Pittsburg, PA) and acute coronal slices (300 µm) were cut on a vibratome (VT1000S, Leica Microsystems; Buffalo Grove, IL) through the striatum in ice-cold, bubbling NMDG based cutting solution. Slices recovered for 15 minutes at 32 °C in bubbling NMDG cutting solution, then transferred to a holding chamber containing normal aCSF [consisting of (in mM): NaCl 125, KCl 2.5, NaH2PO4 1.25, NaHCO3 25, D-Glucose 12.5, MgCl2 1, CaCl2 2, pH = 7.4, 295 mOsm] bubbling (95% O2/5% CO2) at 28 °C. At least one hour after recovery, the slices were placed in the recording chamber, in which normal aCSF (33∼34 °C, bubbling with 95% O2/ 5% CO2) was perfused over the slices at ∼2 mL/min throughout recordings. Dorsal striatal SPNs were visualized under IR-DIC optics (Zeiss Axioskop2; Oberkocken, Germany) at 40x and D1– or D2-SPNs were confirmed by eGFP expression with brief observation in the epifluorescent channel. D1-SPNs (eGFP-positive in D1-eGFP mice, or eGFP-negative in D2-eGFP mice) or D2-SPNs (eGFP-positive in D2-eGFP mice, or eGFP-negative in D1-eGFP mice) that were ChR2-mCherry-negative, but in the injection site and surrounded by cells expressing ChR2-mCherry were targeted for recording. Only animals with high efficiency labeling throughout the cortex were used for recordings to determine collateralization.

Revised version: 4-8 days were allowed for expression and optimal cell health post unilateral (EnvA) SAD-ΔG Rabies-ChR2-mCherry injection before electrophysiology recordings on acute slice were carried out (Klug et al., 2017; Smith et al., 2016). Mice were anesthetized with ketamine/xylazine and transcardially perfused with ∼20 mL ice-cold, bubbling (95% O2/5% CO2) NMDG cutting solution [consisting of (in mM): NMDG 105, HCl 105, KCl 2.5, NaH2PO4 1.2, NaHCO3 26, Glucose 25, Sodium L-Ascorbate 5, Sodium Pyruvate 3, Thiourea 2, MgSO4 10, CaCl2 0.5, 300 mOsm, pH = 7.4]. The extracted brain was blocked coronally with a brain matrix (Zivic Instruments; Pittsburg, PA) and acute coronal slices (300 µm) were cut on a vibratome (VT1000S, Leica Microsystems; Buffalo Grove, IL) through the striatum in ice-cold, bubbling NMDG based cutting solution. Slices recovered for 15 minutes at 32 °C in bubbling NMDG cutting solution, then transferred to a holding chamber containing normal aCSF [consisting of (in mM): NaCl 125, KCl 2.5, NaH2PO4 1.25, NaHCO3 25, D-Glucose 12.5, MgCl2 1, CaCl2 2, pH = 7.4, 295 mOsm] bubbling (95% O2/5% CO2) at 28 °C. At least one hour after recovery, the slices were placed in the recording chamber, in which normal aCSF (33∼34 °C, bubbling with 95% O2/ 5% CO2) was perfused over the slices at ∼2 mL/min throughout recordings. Dorsal striatal SPNs were visualized under IR-DIC optics (Zeiss Axioskop2; Oberkocken, Germany) at 40x and D1- or D2-SPNs were confirmed by eGFP expression with brief observation in the epifluorescent channel. D1-SPNs (eGFP-positive in D1-eGFP mice, or eGFP-negative in D2-eGFP mice) or D2-SPNs (eGFP-positive in D2-eGFP mice, or eGFP-negative in D1-eGFP mice) that were ChR2-mCherry-negative, but in the injection site and surrounded by cells expressing ChR2-mCherry were targeted for recording. This configuration ensured that the distance between recorded and starter cells did not exceed 100 µm, maintaining close anatomical proximity and thereby preserving the likelihood of shared cortical innervation within the examined circuitry. Only animals with high efficiency labeling throughout the cortex were used for recordings to determine collateralization.

Voltage clamp recordings were performed using 3-4 MΩ patch pipettes (WPI; Sarasota, FL), which were pulled from borosilicate glass on a P-97 pipette puller (Sutter Instruments; Novato, CA) and filled with a Cs+ methanesulfonate based internal solution [consisting of (in mM): CsMeSO3 120, NaCl 5, TEA-Cl 10, HEPES 10, QX-314 5 EGTA 1.1, Mg-ATP 4, Na-GTP 0.3, pH = 7.2-7.3, 305 mOsm]. All cells were voltage clamped at -70 mV during recording. Five minutes post break-in, paired light pulses (473nm, 5-25 mW/mm2, 2.5 ms, 50 ms ISI) were delivered through a glass fiber optic (200 µm in diameter, Thor Labs; Newton, NJ), positioned close to the recorded cell (50-150 µm), at 0.05 Hz using a 473 nm blue DPSS laser system (Laserglow Technologies, Toronto, ON). Light evoked currents were collected after at least 8-10 minutes of bath application 50 - 100 µM picrotoxin (MilliporeSigma, St. Louis, MO) to block any ChR2-mediated fast GABAAR transmission. Twenty sweeps were collected to determine latency and CV. At the end of experiments, both 10 µM NBQX (AMPAR antagonist) and 50 µM DL-APV (NMDAR antagonist) (MilliporeSigma, St. Louis, MO) were applied to block AMPAR and NMDAR-mediated transmission, respectively to confirm the EPSCs. Series resistance was initially compensated and monitored continuously throughout the experiment, and the data were rejected if the series resistance changed by more than 20% over the duration of the recording. A cell is considered connected if it has a detectable, reliable current (20 sweeps, 0.05 Hz) with onset latency less than 10 ms post laser-on (Klug et al., 2017; Smith et al., 2016). Voltage clamp recordings were digitized at 10 kHz and filtered at 2 kHz.

For current clamp recordings of rabies-positive pyramidal neurons in the cortex, a potassium methanesulfonate based internal solution [(in mM): KMeSO4 135, KCl 5, CaCl2 0.5, HEPES 5, EGTA 5, Mg-ATP 2, Na-GTP 0.3, (pH = 7.3, 305 mOsm)] was used. 750 ms current injections (-250 to 200 pA) were given to test the membrane potential response of rabies-ChR2 positive pyramidal neurons, in primary motor cortex layer 5, and followed by 20 Hz or 5 Hz optogenetic stimulation to test the response of these neurons to light. Current clamp recordings were filtered and digitized at 10 kHz. All recordings were performed using a Multiclamp 700A amplifier (Molecular Devices; Sunnyvale, CA), digitized with Digidata 1440 (Molecular Devices; Sunnyvale, CA) and collected with pClamp 9 software (Molecular Devices; Sunnyvale, CA). Data were analyzed with Clampfit 9.

Open field

After helper viruses’ injections in the striatum, animals were put back on food and allowed to recover and viral expression. They were then injected with (EnvA) SADΔG-ChR2-mCherry virus in the striatum and implanted with fiber optics in the MCC or M1 as described above (see Surgery and Viral Injection). Then animals were allowed to recover over 3 days. On the fourth day post injection and implantation, animals went through open field test. They were connected to fiber-optic leads (Doric) that connected to a laser through a commutator for free movement. An additional light shield was attached at fiber optic connection to the mouse to mask the laser light output. Following habituation to the fiber optic connections in a home cage the mice were placed in the middle of a 41cm x 41cm square, white and evenly illuminated open field chamber. Custom MEDPC code delivered 20 Hz or 5 Hz stimulation (473 nm blue laser, 5 mW power at connection to mouse, 10 ms pulse width) for 15 seconds after every 3 minutes and 45 seconds, and each animal received 3-4 replicates. Mice with AAV5-EF1 -DIO-ChR2(H134R)-mCherry injected bilaterally in DMS or DLS went through similar open field test after 4 days of recovery from surgery, with optic stimulation in DMS, DLS or M1. Video was collected for each run and analyzed in Ethovision 8.5. To analyze the open field data, the behavior was binned in 10-s bins and distance traveled during laser on period was normalized to the averaged distances during preceding 45 s just prior to stimulation onset.

Optogenetic intracranial self-stimulation (opto-ICSS)

In opto-ICSS experiments, two different subsets of animals were used: to stimulate D1- and D2-SPNs in DMS and DLS, or to stimulate D1- or D2-SPN-projecting cortical neurons in MCC and M1, respectively. Mice that had never experienced the operant chamber were injected with virus and implanted with fiber optics using the procedure described above. From the fourth day following surgery, the mice received ICSS training for 7 consecutive days. They were attached to fiber-optic patch cords and placed in an operant chamber. Each session began with the illumination of a house light and the extension of two levers: one active (left) and one inactive (right). Every time the mouse pressed the active lever, a 20 Hz stimulation was triggered (473 nm blue laser, 5 mW power at connection to mouse, 10 ms pulse width, 1 s duration) targeting the cell bodies in MCC or M1 that project to D1- or D2-SPNs. Each session concludes after 90 minutes with the retraction of the levers and the house light turning off. Continuous pressing of the lever during stimulation will not lengthen the stimulation period. Pressing of the inactive lever had no consequence and was used as a control of general activity measure of non-contingent lever pressing. All protocols were custom written in MEDPC (Med Associates).

Sequence training and optogenetic stimulation

Prior to the injection of the rabies virus, animals were pre-trained for three weeks in fixed ratio 8 (FR8) or fixed ratio 4 (FR4) task (Jin et al., 2010; Jin et al., 2014). Briefly, animals were food-restricted (30 hrs) to start training and weighed daily to monitor their bodyweight. They were fed approximately 2-2.5 g regular chow/mouse/day after each behavioral training session concluded to maintain around 85% of their initial weight. Animals were trained in operant chambers (21.6 cm (L) x 17.8 cm (W) x 12.7 cm (H)) housed in a sound attenuating box (Med-Associates, St. Albans, VT) with two retractable levers to the left and right of a central food magazine and a house light (3 W, 24 V) opposite to the levers and magazine. Sucrose solution (15 µl, 10%) was delivered by a syringe pump into a metal bowl as a reinforcer. Magazine entries were recorded using an infrared beam break detector. Behavioral chambers were controlled by MED-PC IV software (MED Associates, VT) that recorded all timestamps of lever presses and magazine entries with a resolution of 10 ms.

Operant training began with continuous reinforcement (CRF) also known as fixed ratio 1 (FR1) in which animals received a reinforcer following each lever press. The animals were trained on CRF for both levers (separate flanking sessions) over three days and the order of lever presentation was counterbalanced. Each session began with the illumination of the house light and the extension of one lever. The session ended with the offset of the house light and retraction of the lever after 90 minutes of training or after a reinforcer cap was reached. On day 1, 2, 3, the mice could earn up to 10, 15, or 30 sucrose reinforcers, respectively. After the animals acquired CRF over 3 days, they were transitioned to FR4 and FR8 schedules on independent levers and the order counterbalanced. The session began with the illumination of the house light and the extension of either the left or right lever. Following four consecutive lever presses (FR4), mice received a reinforcer in a central magazine port. There was no time requirement for completion of the action sequence. The session concluded with the retraction of the lever and the offset of the house light after the mouse received either 80 reinforcers or 90 minutes expired. Another session was given just following the conclusion of the FR4 session, where eight consecutive lever presses (FR8) on the opposite lever resulted in the delivery of a sucrose reinforcer. The order of training FR4 or FR8 was randomly shuffled over 21 days pre-training. Left and right levers were randomly assigned FR4 or FR8 schedules and that set up was maintained for each animal over pre-training.

On the fourth day after rabies injection and fiber optic implant, and after open field test, the mice were food deprived for 24 hours to start optogenetic test in sequence tasks. On the fifth day, the animals were tethered to two fiber-optic patch cables attached to a commutator (Doric, Canada) to allow for free rotation and placed back in the original pre-training operant box. They were given three days of re-training in a session of FR4 on one lever and subsequent session of FR8 on the opposite lever with fiber attached (90-minute session, 80 reinforcers max). The order of the sessions was randomly shuffled. If the animals successfully completed 80 reinforcers, they were transitioned to optogenetic stimulation test session. On day 8 post rabies injection, optogenetic stimulations (20 Hz, 473 nm blue laser, 5 mW power at connection to mouse, 10 ms pulse width) were randomly delivered for 8 seconds (a time period covering roughly the entire lever press sequence) on the first press (defined by the first lever press after either head entry or 2-second break after the reward delivery) with a 50% likelihood of control non-stimulated trials randomly interleaved (Geddes et al., 2018). Stimulus conditions were repeated on multiple days if needed to collect enough trials for statistics. On day 12 post rabies injection, the animals were perfused for histology analysis.

All sequence data were analyzed in MATLAB using custom scripts. To construct the peri-event time histograms (PETH), all lever presses before the reward (control or stimulation trials) were aligned to the first press of the FR4 or FR8 sequence, averaged in 100 ms bins, and filtered with a Gaussian low-pass filter (window size = 5, standard deviation = 5). All the PETHs were plotted with the first press omitted for illustration and comparison clarity. The effects of optogenetic modulation on press rate were qualitatively similar for FR4 and FR8 sequences and thus combined for statistics.

Histology and microscopy

Approximately twelve days following rabies injection or after behavior tests, mice were anesthetized with an overdose of ketamine/xylazine and transcardially perfused with 0.01 M PBS (30-40 mL) followed by 4% paraformaldehyde (PFA)/0.1 M PB, pH 7.4 (30-40 mL), with a peristaltic perfusion pump (Cole Parmer, Vernon Hills, IL) (Klug et al., 2017; Smith et al., 2016). The brain was carefully extracted and post-fixed in 4% PFA/0.1 M PB overnight (16-24 hrs), then transferred to 30% sucrose/0.1 M PB for 1-2 days until the brain equilibrated and sunk. On the day of cutting, it was coronally blocked with a brain matrix (Zivic Instruments; Pittsburg, PA) and mounted on a freezing microtome. Coronal slices were collected from the most rostral to the most caudal sites at 50 µm resolution in 96 well plates containing cryoprotectant (0.1 M phosphate buffer, ethylene glycol, glycerol) to maintain AP position. Brain slices surrounding the injection site and fiber implant site were mounted on super frost plus slides (Thermo Fisher Scientific, Waltham, MA), counterstained for DAPI and cover slipped with Aqua-Poly/Mount mounting media (Polysciences, Inc; Warrington, PA). Slides were scanned on an automated slide scanner (Olympus VS120) at 10x in the blue and red channels. Images were batch converted to composite TIFFs and saved for image analysis.

Statistics

Statistics were conducted in Graph Pad Prism 6.01 (La Jolla, CA). Fisher’s exact-test was used in comparing the likelihood of connections in slice recordings. Student unpaired two-tailed t-test was used in open field test and sequence operant task to analyze optogenetic stimulation effects. Non-parametric Mann Whitney U Test was conducted when distributions significantly deviated from normal distributions. The reinforcing effect of each group in the ICSS task was assessed using a permutation test (10,000 permutations per group). A one-sample two-tailed t-test or Wilcoxon signed-rank test was used to assess the effectiveness of optogenetic activation in each individual group during the sequence operant task.